Kurs 5.3 "Vom Gehirn Lernen" -- Python-Tutorial  | [Startseite](index.ipynb)

---

# 101 - Grundlagen der Programmierung in Python

- [Aufgabe 1 - Werte und Typen](#Aufgabe-1---Werte-und-Typen)
- [Strings](#Strings)
- [Aufgabe 2 - Strings und Dokumentation](#Aufgabe-2---Strings-und-Dokumentation)
- [Aufgabe 3 - String Formatierung](#Aufgabe-3---String-Formatierung)
- [Listen und Tupel](#Listen-und-Tupel)
- [Slicing](#Slicing)
- [Dictionaries](#Dictionaries)
- [Aufgabe 4 - Dictionaries](#Aufgabe-4---Dictionaries)

## Über Variablen

Pythoncode besteht aus einer Aneinanderreihung von *Befehlen*, die der sogenannte Python-Interpreter _nacheinander_ ausführt. Mit anderen Worten geht der Interpreter Zeile für Zeile durch ein Python-Script oder, wie in unserem Fall, durch eine Jupyter-Codezelle und übersetzt das, was in der jeweiligen Zeile steht, in elementare "Operationen", die dann auf der CPU ausgeführt werden können.

Ein solcher Befehl ist die Zuweisung eines Wertes zu einer **Variablen**. Dabei übersetzt der Interpreter den Wert der Variablen in Nullen und Einsen (wir nennen das _binäre Darstellung_) und schreibt diese an einen Ort im Arbeitsspeicher.
Eine mögliche Übersetzung von der natürlichen Zahl 42 in das Binärsystem ist `0010 1010`, die Codierung des Buchstabens "A" ist `0010 0001` (nach dem [ACSII-Standard](https://de.wikipedia.org/wiki/American_Standard_Code_for_Information_Interchange)).
Und gleichzeitig bekommt der Ort einen Namen, den wir uns Menschen gut merken können, damit wir später wieder auf diese Variable zugreifen können.


Im unteren Beispiel weisen wir der Variablen mit dem Namen `x` den Wert `42` zu.

```python
x = 42
# Diese Zelle auswählen und mit `<SHIFT> + <ENTER>` ausführen.
# Der letzte Rückgabewert der Zelle wird dann unten ausgegeben,
# hier also der Wert von `x`:
x
```

> **Hinweis:** Dieses Kursmaterial besteht aus interaktiven [Jupyter Notebooks](http://jupyter.org). Jede _Zelle_, wie die obige, könnt ihr per Klick oder mit den Pfeiltasten auswählen und mit `<SHIFT> + <ENTER>` ausführen. Probiert's mal mit der Zelle oben!
> 
> **Geht, um dem Kurs zu folgen, _alle_ Zellen nacheinander durch und führt sie aus.**
>
> Wie hier demonstriert gibt das Jupyter Notebook immer den letzten Rückgabewert einer Zelle aus. 

Wir können auch die `print()` Funktion verwenden, um den Wert einer Variablen auf dem Bildschirm auszugeben.

```python
print(x)
```

(Guter) Code besteht nicht nur aus Befehlen, sondern auch aus Kommentaren. Das ist Text in deinem Code, der vom Interpreter übersprungen wird. Das ist ganz nützlich, um z.B. in menschlicher Sprache zu beschreiben, was im Code passiert. Oder einfach, um zu Testzwecken Codezeilen zu deaktivieren.

In Python werden Kommentare mit einer Raute (`#`) gekennzeichnet. Alles, was in der Zeile nach dem `#` kommt, wird vom Interpreter übersprungen.

```python
# Ein Kommentar
print("Kein Kommentar!")  # wieder ein Kommentar
# print('Diese Zeile wird nicht interpretiert, obwohl sie ein regulärer Pythonbefehl wäre')
```

Damit der Interpreter weiß, wie er einen Wert in Nullen und Einsen übersetzen muss (oder umgekehrt, wie er eine bestimmte Kombination von Nullen und Einsen interpretieren soll), haben Variablen immer einen sogenannten **Datentyp**. 
Häufige Datentypen sind:

- `int` für ganze Zahlen (englisch: _integer_), z.B. `1`, `42`, `-10`, 
- `float` für Fließkommazahlen, z.B. `0.5`, `3.14`, `1e10`,
- `str` für Zeichenketten (_Strings_), z.B. `"Hello World!"`,
- `boolean` für Wahrheitswerte (_Booleans_), also `True` und `False`

> Die Typen der Python Standardlibrary findet ihr in der [Dokumentation](https://docs.python.org/3/library/stdtypes.html).

**Gut zu wissen:** In Python (im Gegensatz zu manch anderen Programmiersprachen wie z.B. C) sind Variablen dynamisch typisiert. Das heißt, dass der Interpreter automatisch erkennt, welche Datentypen man verwendet und den Typ einer Variablen entsprechend anpasst.

**Hinweis:** Mit der Funktion `type` kann man sich ausgeben lassen, welcher Datentyp eine Variable hat

```python
type(x)
```

```python
x = 0.5  # jetzt ändern wir den Wert von x zu 0.5, und x wird damit zu einem float
type(x)
```

Weiterhin können Werte in einen anderen Typ konvertiert werden:

```python
x = int(0.5)  # Bei der Konvertierung zu int wird immer abgerundet!
print(x)

y = 2
print(type(y))
y = float(y)
print(y, type(y))
```

Die grundlegenden Rechenoperationen `+`, `-`, `*`, `/` und `**` (= Potenz) sind ebenfalls in Python verfügbar und verhalten sich, wie man es erwartet:

```python
1 + 3
```

```python
3 / 2
```

**Gut zu wissen:** Das Ergebnis einer Rechenoperation bekommt automatisch den "richtigen" Datentyp zugewiesen. Z.B. resultiert eine Division zweier Integers in einem Float oder die Addition eines Floats und eines Integers ebenso in einem Float:

```python
3.14 + 6
```

```python
3**2
```

### Aufgabe 1 - Werte und Typen

Der Operator zur Potenzierung ist in Python `**`. Weise einer Variablen `y` folgende Werte zu und lasse dir dann ihren Wert und Typ ausgeben. Stimmen Wert und Typ mit deinen Erwartungen überein?

a) $4^3$

> **Hinweis:** Dies ist eine Übungsaufgabe. Verwende die Zelle unten, um sie zu lösen. **Entferne** dazu den Platzhalter-Code
>
> ```python
> # DEINE LÖSUNG HIER
> ```
>
> und schreibe stattdessen den Code zur Lösung der Aufgabe.

<!-- grade_id: 101-1a-sol, points: 0, grade: false, solution: true -->
```python
# DEINE LÖSUNG HIER
```

> **Hinweis:** Auf die Lösung jeder Aufgabe folgt eine Test-Zelle wie die folgende, mit der du deine Lösung überprüfen kannst. Führe einfach diese Zelle aus. Wenn etwas an deiner Lösung nicht stimmt, erhälst du eine Fehlermeldung mit Hinweisen.

```python
# DEINE LÖSUNG HIER
```

<!-- grade_id: 101-1a-test, points: 1, grade: true, solution: false -->
```python
try:
    y
except NameError:
    raise NameError(
        "Es gibt keine Variable 'y'. Weise den Wert einer Variablen mit diesem Namen zu."
    )

assert y == 64, "Die Variable hat nicht den richtigen Wert. Überprüfe deine Rechnung."
print("👍 Sehr gut!")
```

b) $2+3.4^2$

<!-- grade_id: 101-1b-sol, points: 0, grade: false, solution: true -->
```python
# DEINE LÖSUNG HIER
```

<!-- grade_id: 101-1b-test, points: 1, grade: true, solution: false -->
```python
try:
    y
except NameError:
    raise NameError(
        "Es gibt keine Variable 'y'. Weise den Wert einer Variablen mit diesem Namen zu."
    )

assert (
    y == 2 + 3.4**2
), "Die Variable hat nicht den richtigen Wert. Überprüfe deine Rechnung."
print("Ok, du hast es verstanden.")
```

## Strings

Strings sind **Zeichenketten** wie:

```python
s = "Hello World"
```

und werden in Python vom Typ `str` repräsentiert. Um einen String zu erstellen können wir einzelne (`'`), doppelte (`"`) oder dreifache (`'''` oder `"""`, für mehrzeilige Strings) Anführungszeichen verwenden, sodass das jeweils andere Anführungszeichen im String verwendet werden kann:

```python
"I'm happy!"
```

Alternativ können Steuerzeichen im String auch _escaped_ werden:

```python
'Say "hi"'
```

### Strings sind Reihen

Da Strings eine Aneinanderreihung von Elementen (in diesem Fall Textzeichen) darstellen, können viele Funktionen mit Strings verwendet werden, die mit **Reihen** arbeiten. Dazu gehören:

```python
len(s)  # gibt die Zahl der Reihenelemente zurück
```

```python
s[0]  # gibt das Element der Reihe an der Stelle 0 zurück
```

```python
s + "!"  # Reihen können kombiniert werden
```

Wir werden weiter unten bei den [Listen](#-Listen-und-Tupel) noch mehr über Reihen erfahren.

### Strings sind Objekte

Die meisten "Dinge" in Python sind **Objekte**, d.h. neben ihrem Typ besitzen sie assoziierte **Attribute** und **Methoden**, auf die über einen Punkt `.` zugegriffen werden kann.
Neben Strings sind bspw. auch Werte der anderen schon bekannten Datentypen wie `int(5)` und sogar Funktionen und die Datentypen selbst Objekte.

Objekte sind die zentralen Bestandteile des _objektorientierten Programmieren_ (kurz OOP).
Eine andere Art, Programme zu designen, ist das _funktionale Programmieren_.
Was Funktionen genau sind, werden wir später noch diskutieren (aber ein paar kennst du schon: z.B. `type()`, `print()` oder `len()`).

Über OOP allein könnte man wahrscheinlich etliche DSA-Kurse abhalten, daher behandeln wir sie hier nur ganz kurz.
Im Wesentlichen ist ein Objekt eine Art Variable, die ihrerseits aus "versteckten" Variablen zusammengesetzt ist.
Diese "versteckten" Variablen nennen wir _Attribute_.
_Methoden_ wiederum sind Befehle, die etwas mit einem Objekt machen können, d.h. sie können auf die Attribute zugreifen, diese auslesen, sie verändern und mit ihnen rechnen.
Die Syntax für Methoden ist: `object.method()`, d.h. die Methode wirkt auf das Objekt vor dem Punkt.

Eine _Klasse_ wiederum fasst alle Objekte einer bestimmten Sorte zusammen und definiert damit, welche Attribute ein Objekt hat und welche Methoden darauf angewendet werden können.

Alles klar so weit?
Ok, das war ein bisschen viel auf einmal! Deshalb ein anschauliches Beispiel:

Nehmen wir die Klasse `Hund`.
Ein Hund könnte z.B. die Attribute `name`, `fellfarbe` und `x`- und `y` Koordinaten (Hunde rennen ja bekanntermaßen viel, und daher ist es ganz praktisch zu wissen, wo sie sind...) haben. Eine mögliche Methode wäre z.b. `.lauf()`.

In Python könnte das dann so aussehen (**Hinweis:** Da wir keine echte Hundeklasse implementiert haben, funktioniert die unsere Zelle natürlich nicht!)

```python
lumpi = Hund(
    name="Lumpi", fellfarbe="schwarz", x=3.2, y=4.5
)  # Initialisiere den Hund lumpi
# lumpi ist jetzt ein Objekt der Klasse Hund
print(lumpi.fellfarbe)  # 'schwarz'
lumpi.lauf(4.7, 2.1)
print(lumpi.x, lumpi.y)  # '7.9, 6.6'
```

Im Moment müsst ihr noch nicht genau verstehen, wie Objekte im Python funktionieren. 
Es ist nur ganz hilfreich zu wissen, _dass_ es Objekte in Python gibt, und wie wir Methoden darauf anwenden können.

---
Wie schon gesagt: 
Strings sind Objekte und haben damit interne Attribute und Methoden. 
Einige nützliche Methoden sind z.B.

```python
s.upper()
```

```python
s.split()
```

```python
s.index("World")
```

> **Hinweis:** In Jupyter Notebooks können wir die **`<TAB>`-Vervollständigung** verwenden um die assoziierten Attribute und Methoden eines Objekts zu inspizieren:
>
> ```python
> s = "Hello World"
> # Zelle ausführen, dann:
> s.<TAB> # `s.` tippen und die <TAB>-Taste drücken
> ```
>
> Dies zeigt die verfügbaren Attribute und Methoden des Strings `s` an. Die `<TAB>`-Vervollständigung für eine Variable funktioniert erst nachdem die Variable erstellt wurde, also die Zelle einmal ausgeführt wurde.
>
> Um herauszufinden, was eine Funktion oder Methode tut, könnt ihr im Jupyter Notebook ein Fragezeichen `?` verwenden:
>
>```python
>In [1]: s.split?
>```
>```markdown
>Docstring:
>S.split(sep=None, maxsplit=-1) -> list of strings
>
>Return a list of the words in S, using sep as the
>delimiter string.  If maxsplit is given, at most maxsplit
>splits are done. If sep is not specified or is None, any
>whitespace string is a separator and empty strings are
>removed from the result.
>Type:      builtin_function_or_method
>```
>
> Schreibt ihr stattdessen zwei Fragezeichen `??` zeigt das Jupyter Notebook die gesamte Definition der Funktion oder Methode an.
>
> **Verwendet die `<TAB>`-Vervollständigung und die `?`-Dokumentation häufig um hilfreiche Attribute und Methoden zu finden und zu verstehen!**

### Aufgabe 2 - Strings und Dokumentation

Finde im folgenden String mithilfe einer Methode, wie häufig der Buchstabe `"A"` auftaucht und weise den Wert einer Variable `n` zu. Probiere die `<TAB>`-Vervollständigung und die `?`-Dokumentation aus um eine passende Methode zu finden.

<!-- grade_id: 101-2-sol, points: 0, grade: false, solution: true -->
```python
s = "CAGTACCAAGTGAAAGAT"
### BEGIN SOLUTION
n =  # tippe s. und dann <TAB>, um nach einer passenden Methode zu suchen  # noqa
### END SOLUTION
print(n)
```

<!-- grade_id: 101-2-test, points: 1, grade: true, solution: false -->
```python
try:
    y
except NameError:
    raise NameError(
        "Es gibt keine Variable 'n'. Weise den Wert einer Variablen mit diesem Namen zu."
    )
assert (
    n == 8
), "Das ist nicht die richtige Anzahl. Versuch's mal mit der `count` Methode!"
print("Klappt!")
```

### String-Formatierung

Eine wichtige Methode ist `str.format()`, die markierte Platzhalter im String mit Werten ersetzt, die der Methode übergeben werden:

```python
x = 10
y = 3.1415
print("Der Wert von x ist {} und Pi ist ungefähr {}".format(x, y))

name1 = "Anna"
name2 = "Bob"
print("Hallo {} und {}!".format(name1, name2))
```

> **Hinweis:** Außerdem kann man Zahlen formatieren. Wie das funktioniert, siehst Du in der offiziellen [Python-Dokumentation]() oder [hier](https://pyformat.info/).

> **Noch ein Hinweis:** Seit Python 3.6 gibt es noch eine Variante, die viel besser lesbar ist: 
Die sogenannten ["f-strings"](https://docs.python.org/3/reference/lexical_analysis.html#f-strings). 
Wie sie im Vergleich zu `.format()` funktioniert, wird in diesem [Tutorial auf realpython.com](https://realpython.com/python-string-formatting/#3-string-interpolation-f-strings-python-36) erklärt.

### Aufgabe 3 - String Formatierung

Schreibe deinen Namen in die Variable `name`. Verwende `.format()` um aus `s` und `name` den Gruß `"Hello World, mein Name ist __DEIN_NAME__!"` zusammenzusetzen. Weise den zusammengesetzten String der Variable `greeting` zu.

<!-- grade_id: 101-3-sol, points: 0, grade: false, solution: true -->
```python
s = "Hello World"
name = "__DEIN_NAME__"
### BEGIN SOLUTION
greeting = s +  # schreibe hier einen f-string
### END SOLUTION
print(greeting)
```

<!-- grade_id: 101-3-test, points: 1, grade: true, solution: false -->
```python
try:
    greeting
except NameError:
    raise NameError(
        "Es gibt keine Variable 'greeting'. Weise den Wert einer Variablen mit diesem Namen zu."
    )
assert "mein Name ist" in greeting
print("Hello {}! 👋".format(name))
```

## Listen und Tupel

Während Strings einzelne Zeichen aneinanderreihen, repräsentieren **Listen** und **Tupel** eine Reihe _beliebiger_ Werte. 
Die Elemente einer Liste können verändert werden, während ein Tupel unveränderlich ist. 
Wir können Listen erstellen, indem wir die Elemente durch Kommata getrennt in eckigen Klammern `[` und `]` angeben, und Tupel durch die Verwendung runder Klammern `(` und `)`:

```python
l = [4, 0.5, "Alice"]
l
```

```python
t = ("Bob", True)
t
```

### Auch Listen und Tupel sind Reihen

Wie bei jeder Reihe können wir die Anzahl der Elemente bestimmen und auf einzelne Elemente über ihren Index zugreifen:

```python
len(l), len(t)
```

```python
l[0]  # Indexierung beginnt in Python mit dem Index 0
```

```python
l[1]
```

```python
l[2]
```

```python
t[0]
```

```python
t[1]
```

### Listen sind veränderlich

Anders als Strings und Tupel können Listen jedoch verändert werden, indem Elemente verändert, hinzugefügt oder entfernt werden:

```python
l[1] = -2.2  # Weise der Liste einen neuen Wert beim Index 1 zu
l
```

```python
l.append(-3)  # Füge einen neuen Wert am Ende der liste hinzu
l
```

```python
l.pop()  # Entfernt das letzte Element (bzw. das am angegebenen Index) aus der Liste und gibt es zurück
l
```

## Slicing

Du kannst mit der _Slicing_ Syntax auf Teile einer Reihe zugreifen:

```python
slice = list[start:stop:step]
```

Dabei bezeichnet `start` den ersten Index und `stop` den letzten nicht mehr enthaltenen Index des Abschnitts. Mit `step` kannst du eine Schrittweite angeben, in der die Reihe iteriert werden soll.

Du musst nicht alle drei Argumente angeben. Per default ist dann `start=0`, `stop=len(list)` und `step=1`:

```python
l[:2]  # Der erste Index ist per default start=0
```

```python
l[2:]  # Der letzte Index ist per default das Ende der Reihe
```

```python
l[::2]  # Jedes zweite Element
```

```python
l[::-1]  # Umgekehrte Reihenfolge
```

> **Hinweis:** Slicing ist ein mächtiges Werkzeug um kurzen, prägnanten Code zu schreiben und wird dir bei der Arbeit mit _Numpy Arrays_ ab Lektion [201 - Numerik mit Numpy](201 - Numerik mit Numpy.ipynb) noch sehr häufig begegnen. TODO: evtl. Link anpassen!

## Dictionaries

Ein weiterer wichtiger Datentyp ist das **Dictionary**. Ein Dictionary ordnet jeweils einem _Key_ einen _Value_ zu und wird durch geschweifte Klammern `{` und `}` erstellt:

```python
d = {"a": 1, "b": 2, "c": 3}
d
```

Auf einzelne Werte kann über ihren Key zugegriffen werden:

```python
d["a"]
```

Dictionaries sind veränderlich wie Listen:

```python
d["d"] = 4
d
```

Es müssen nicht unbedingt Strings als Keys verwendet werden, und auch verschiedene Datentypen sind möglich:

```python
e = {"some_key": 4.2, 3517: "some_value"}
e
```

```python
e["some_key"]
```

```python
e[3517]
```

Nützliche Methoden für dictionaries:

```python
d.keys()  # gibt alle keys zurück
```

```python
d.values()  # gibt alle items zurück
```

```python
d.items()  # gibt alle key-value-Paare als Tupels zurück
```

```python
len(d)  # auch Dictionaries haben eine Länge!
```

## Aufgabe 4 - Dictionaries

Schreibe ein Wörterbuch mit einigen deutschen Wörtern und ihrer Übersetzung ins Englische. Weise es der Variable `d` zu und verwende es dann, indem du dir einige Werte darin ausgeben lässt.

<!-- grade_id: 101-4-sol, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
d = {"Wörterbuch": "dictionary"}
### END SOLUTION
d["Wörterbuch"]
```

<!-- grade_id: 101-4-test, points: 1, grade: true, solution: false -->
```python
try:
    d
except NameError:
    raise NameError(
        "Es gibt keine Variable 'd'. Weise den Wert einer Variablen mit diesem Namen zu."
    )
assert len(d) > 1, "Das Wörterbuch hat zu wenig Einträge."
print("Wow, {} Einträge im Wörterbuch!".format(len(d)))
```

---

Du kannst nun mit Variablen und Datentypen umgehen und den Python Interpreter rechnen lassen. 
Lerne in der nächsten Lektion, wie du mit _Control Flow_ Anweisungen deine ersten Programme schreiben kannst.

> **Tipp**: Erstelle als kleine Übung dein eigenes "Python-Cheatsheat", in dem du die wichtigsten Python-Befehle für dich zusammenfasst. 
Es gibt zwar bestimmt schon hunderte Cheatcheats fertig im Internet zum Herunterladen. 
Aber du lernst viel mehr dabei, wenn du dir selber eins erstellst! Und du kannst es dann in den nächsten Lektionen noch weiter ausbauen! :)

[Startseite](index.ipynb) | [**>> 102 - Control Flow**](102_Controlflow.ipynb)