Kurs 5.3 "Vom Gehirn Lernen" -- Python-Tutorial | [Startseite](index.ipynb)

---

# 103 - Funktionen und Module

- [Funktionen sind die Bausteine eines Programms](#Funktionen-sind-die-Bausteine-eines-Programms)
- [Aufgabe 1 - Mehr Primzahlen](#Aufgabe-1---Mehr-Primzahlen)
- [Aufgabe 2 - Fakultät](#Aufgabe-2---Fakultät)
- [Funktionen aus Modulen importieren](#Funktionen-aus-Modulen-importieren)
- [Aufgabe 3 - Globale Variablen](#Aufgabe-3---Globale-Variablen)
- [Aufgabe 4 - Noch mehr globale Variablen](#Aufgabe-4---Noch-mehr-globale-Variablen)
- [Aufgabe 5 - Gaussian](#Aufgabe-5---Gaussian)

Wir können nun bereits Variablen setzen und _control flow_ Anweisungen geben um kleine Programme zu schreiben. Damit können wir schon vieles berechnen! Wenn wir jedoch komplexere Berechnungen ausführen möchten, kann ein Programm schnell sehr lang und unleserlich werden. Eine wichtige Regel in der Programmierung ist es daher, **Wiederholung zu vermeiden**. Dazu definieren wir wiederverwendbare **Funktionen** und verwenden schon vorhandene Funktionalität aus **Modulen**.

## Funktionen sind die Bausteine eines Programms

Eine Funktion ist ein Codeblock, der eine abgeschlossene Aufgabe erfüllt. Funktionen haben immer einen **Namen**, können **Argumente** annehmen und **Rückgabewerte** zurückgeben. In Python haben Funktionen folgende Syntax:

```python
def function_name(arguments):
    # code here
    return values
```

Beachtet wieder die Abgrenzung des Codeblocks durch Einrückung, wie wir es bereits bei `if`-Abfragen und Schleifen kennengelernt haben.

Ist die Funktion definiert, können wir sie mit folgender Syntax aufrufen:

```python
function_name(arguments)
```

Mit Funktionen können wir ein komplexes Problem in lösbare Teilprobleme zerlegen, die wir dann zu einem vollständigen Programm zusammensetzen.

> **Beispiel:** Haben wir einmal eine Funktion geschrieben, die eine Liste sortiert, können wir immer darauf zurückgreifen, anstatt den Code jedes mal aufs Neue zu schreiben.

> **Hinweis:** Nur weil du Code in Funktionen auslagern _kannst_ solltest du das nicht immer tun. Schreibe dann eine Funktion, wenn du dadurch Wiederholungen vermeidest oder das Programm klarer strukturierst. **Häufig ist eine Funktion dann sinnvoll, wenn du ihr einen deskriptiven Namen geben kannst.**

### Argumente und Rückgabewerte

Eine Funktion kann mehrere Argumente annehmen...

```python
def add(a, b):
    summe = a + b
    return summe


print(add(1, 3))
print(add(1.0, 3.2))
print(add(4, 3.0))
```

Man kann diese Funktion auch etwas kürzer schreiben, indem man die Summe gleich an der Stelle berechnet, an der die Funktion ihr Ergebnis zurück gibt:

```python
def add(a, b):
    return a + b


print(add(1, 3))
print(add(1.0, 3.2))
print(add(4, 3.0))
```

Eine Funktion kann auch mehrere Werte zurückgeben:

```python
def double_and_halve(value):
    return value * 2, value / 2


print(double_and_halve(5))
```

Die Rückgabewerte können wir einer oder mehreren Variablen zuweisen:

```python
d, h = double_and_halve(5.0)
print(d)
print(h)
```

Funktionen können **andere Funktionen aufrufen**:

```python
def do_a():
    print("doing A")


def do_b():
    print("doing B")


def do_a_and_b():
    do_a()
    do_b()


do_a_and_b()
```

Argumente können auch einen **_default_-Wert** besitzen und damit **optional** sein:

```python
def say_hello(to_name="World"):
    print("Hello {}!".format(to_name))


say_hello()
say_hello("Alice")
```

Argumente können in der **Reihenfolge** gegeben werden, in der die Funktion sie definiert, oder mit Angabe des Argumentnamens in beliebiger Reihenfolge:

```python
def say_hello(to_name="World", my_name=None):
    if my_name is None:
        print("Hello {}!".format(name))
    else:
        print("Hello {}! My name is {}.".format(to_name, my_name))


say_hello("Alice", "Bob")
say_hello(my_name="Bob", to_name="Alice")
```

### Aufgabe 1 - Mehr Primzahlen

Definiere eine Funktion mit dem Namen `is_prime`, die eine Zahl annimmt und `True` zurückgibt wenn die Zahl eine Primzahl ist, bzw. `False` wenn nicht.

Kopiere dazu deinen Code um Primzahlen zu finden von zuvor und verändere ihn entsprechend.

<!-- grade_id: 103-1-sol, points: 0, grade: false, solution: true -->
```python
def is_prime(n):
    # DEINE LÖSUNG HIER
    pass  # `pass` ist ein Platzhalter und sagt Python, dass es einfach nix machen soll. Entferne es, wenn du eine Lösung hast
```

<!-- grade_id: 103-1-test, points: 1, grade: true, solution: false -->
```python
try:
    is_prime
except NameError:
    raise NameError(
        "Es gibt keine Funktion mit dem Namen 'is_prime'. Stelle sicher, dass deine Funktion so benannt ist."
    )

assert is_prime(
    47
), "Die Zahl 47 sollte eine Primzahl sein. Prüfe den Code deiner Funktion."
assert not is_prime(
    48
), "Die Zahl 48 sollte _keine_ Primzahl sein. Prüfe den Code deiner Funktion."
print("Klappt.")
```

### Aufgabe 2 - Fakultät

Schreibe eine Funktion mit dem Namen `factorial`, die eine Zahl annimmt und die Fakultät `n! = n*(n-1)*...*3*2*1` dieser Zahl zurückgibt.

**Hinweis:** Du kannst zunächst versuchen, die Aufgabe mit einer Schleife zu lösen. Funktionen können sich jedoch auch selbst aufrufen (sog. _rekursive_ Funktionen). Versuche eine Funktion zu schreiben, die _keine Schleife_ verwendet!

<!-- grade_id: 103-2-sol, points: 0, grade: false, solution: true -->
```python
def factorial(n):
    # DEINE LÖSUNG HIER
    pass  # `pass` ist ein Platzhalter und sagt Python, dass es einfach nix machen soll. Entferne es, wenn du eine Lösung hast
```

<!-- grade_id: 103-2-test, points: 1, grade: true, solution: false -->
```python
try:
    factorial
except NameError:
    raise NameError(
        "Es gibt keine Funktion mit dem Namen 'factorial'. Stelle sicher, dass deine Funktion so benannt ist."
    )

assert factorial(1) == 1, "1! sollte 1 ergeben. Prüfe den Code deiner Funktion."
assert factorial(2) == 2, "2! sollte 2 ergeben. Prüfe den Code deiner Funktion."
assert factorial(5) == 120, "5! sollte 120 ergeben. Prüfe den Code deiner Funktion."
print("🙌 Funktioniert!")
```

### Variablen sind dort verfügbar, wo sie definiert wurden

Argumente von Funktionen und Variablen, die im Codeblock der Funktion definiert wurden, sind nur innerhalb der Funktion verfügbar (**_local scope_**):

```python
def do_something():
    local_var = 17  # Diese Variable ist innerhalb der Funktion definiert...
    print(
        "Innerhalb der Funktion ist local_var bekannt und ihr Wert kann hier geprinted werden:"
    )
    print(local_var)


do_something()

print("...aber ausserhalb der Funktion nicht")
print(local_var)
```

Variablen, die zum Zeitpunkt außerhalb von Funktionen definiert wurden, werden **globale Variablen** genannt und sind auch innerhalb von Funktionen verfügbar (**_global scope_**):

```python
PI = 3.14  # Diese Variable ist global definiert...
# ...und kann überall, also auch innerhalb von verschiedenen Funktionen verwendet werden:


def degrees_to_radians(degrees):
    return degrees / 180 * PI


degrees_to_radians(90)
```

**Achtung:** Hier ist die Reihenfolge der Definitionen wichtig! Eine globale Variable die in einer Funktion benutzt werden soll, muss bevor die Funktion aufgerufen wird (*nicht bevor die Funktion definiert wird*) definiert werden.

```python
# Location b (für Aufgabe 3)

GLOBAL_VAR = 12  # Hier definieren wir die globale Variable
print("Hier wird GLOBAL_VAR ausserhalb der Funktion geprinted:")
print(GLOBAL_VAR)


# Hier definieren wir eine Funktion die die globale Variable benutzen wird sobald die Funktion ausgeführt wird
def print_global_var():
    print("Hier wird GLOBAL_VAR aus der Funktion heraus geprinted:")
    print(GLOBAL_VAR)


# Location e

GLOBAL_VAR = 17  # Hier aendern wir den Wert der globalen Variable (um ein bisschen Verwirrung zu stifen)

# Location d

# Und jetz erst führen wir die Funktion aus
print_global_var()

# Location c
```

### Aufgabe 3 - Globale Variablen

**a)** Überlege dir, ob du genau erklären kannst, warum welcher Wert von `GLOBAL_VAR` wo geprinted wird.

**b)** Überlege dir, was passieren würde, wenn wir die Funktionsdefinition  `def print_global_var():` ganz nach oben in der Zelle verschieben würden (Location b).

**c)** Überlege dir, was passieren würde, wenn wir die Funktionsdefinition  `def print_global_var():` ganz nach unten in der Zelle verschieben würden (Location c).

**d)** Überlege dir, was passieren würde, wenn wir die Funktionsdefinition  `def print_global_var():` nach `GLOBAL_VAR = 17` verschieben würden (Location d).

**e)** Überlege dir, was passieren würde, wenn wir den Funktionsaufruf `print_global_var()` vor `GLOBAL_VAR = 17` verschieben würden (Location e).

**Wenn du deine Antworten zu den Teilaufgaben b) - e) weiter unten überprüfst, ist es wichtig, dass du JEDES MAL bevor du die Zelle ausführst folgende Schritte ausführst:**
- Klicke oben in der Menü-Leiste auf das Tab mit dem Namen "Kernel" (5. von links)
- Klicke dann im Sub-Menü auf "Restart Kernel..."

Der Grund dafür ist, dass sich Jupyter Notebooks ein bisschen anders verhalten als normale Python Skripte. Wenn ein Python Skript einmal ausgeführt wird und dann an seinem Ende ankommt, wird es geschlossen und alle Variablen und deren Werte etc werden vergessen. Von einer Ausführung des Skripts zur nächsten wird nichts übertragen. Bei einer Jupyter Notebook Zelle ist das anders: Die Variablen und ihre Werte werden nicht vergessen wenn eine Zelle zu Ende gelaufen ist und dann anschließend nochmal gestartet wird. 

In dieser Aufgabe soll das Verhalten von globalen Variablen in einer "normalen" Python-Benutzung (also nicht im Jupyter Sonderfall) angeschaut werden, deswegen müssen wir Jupyter zwingen Dinge wieder zu vergessen. Und das machen wir durch den Neustart des Kernels.

```python
# Hier kannst du deine Antworten von oben überprüfen. DENKE DRAN DEN KERNEL JEDES MAL NEU ZU STARTEN (siehe oben)

# MODIFIZIERTER CODE HIER
```

### Lokale Variablen werden bevorzugt

Wenn eine Variable sowohl global als auch lokal definiert ist, wird die lokale Variable bevorzugt:

```python
a = 1


def show_var():
    a = 2
    print(a)


show_var()  # Hier wird die lokale Variable verwendet...
print(a)  # ... und hier die Globale.
```

### Achtung!

Bei der Verwendung von globalen Variablen ist Vorsicht geboten. Je länger der Code, desto schwieriger ist es den kompletten Überblick zu behalten, was welche Variable wo tut und wo sie wie verändert wird. 

### Aufgabe 4 - Noch mehr globale Variablen

Versuche, ohne den untenstehenden Code laufen zu lassen, vorherzusagen was nacheinander geprinted wird.

```python
a = 42
b = 5

def function1():
    b = 17
    print('a = {}'.format(a))
    print('b = {}'.format(b))
    
def function2():
    a = 42
    print('a = {}'.format(a))
    print('b = {}'.format(b))

function2()    
    
print('a = {}'.format(a))
print('b = {}'.format(b))

a = 13

function1()

print('a = {}'.format(a))
print('b = {}'.format(b))

b = 13
a = 17

function1()
function2()    
```

```python

```

**Um genau diese potentielle Verwirrung zu vermeiden, werden globale Variablen so selten wie möglich genutzt!** Wenn die Funktion Input-Parameter benötigt solltest du sie der Funktion immer als Argumente übergeben. Wenn überhaupt, werden globale Variablen für Werte verwendet, die sich nie während des Programmablaufs ändern werden, z.B. physikalische oder mathematische Konstanten wie $\pi$.

> Per Konvention schreiben wir Konstanten die wir als globale Variablen programmieren in Großbuchstaben wie `PI`.

## Funktionen aus Modulen importieren

Die beruhigende Nachricht ist: viele Probleme wurden schon gelöst. Für häufige Aufgaben, wie bspw. das Sortieren von Listen, existieren sogar hochoptimierte und getestete Lösungen, die wir tunlichst verwenden sollten, anstatt unsere eigene zu schreiben!

Abgesehen von einigen grundlegenden Datentypen und Funktionen wie `print` oder `len` sind diese Funktionen nicht in der Python Standard Library enthalten sondern in **Modulen** ausgelagert. Mit der folgenden Syntax können wir ein Modul **importieren**, um auf die enthaltenen Funktionen zugreifen zu können:

```python
import module
module.function_name()
```

Häufig verwendeten Modulen können wir einen abgekürzten Namen geben:

```python
import module as m
m.function_name()
```

> Hinweis: Manchmal enthalten zwei unterschiedliche Module Funktionen mit dem gleichen Namen, wie z.B. math.exp() und np.exp(). Damit Python weiß, welche Implementierung benutzt werden soll, müssen wir das Modul mit der Punktnotation spezifizieren. Dieses Konzept wird auch Namespaces genannt.

> Achtung: Es hält dich übrigens niemand davon ab eine eigene Funktion zu schreiben, die z.B. `len()` heißt...Du "überschreibst" damit aber dann eine schon existierende Python Funktion! Pass also auf, wie du deine Funktionen oder Variablen benennst! Der Editor (oder Jupyter) können dir dabei helfen: Standard-Pythonbefehle werden meistens anders gefärbt dargestellt. Wenn du also versehentlich eine Funktion oder Variable wie ein schon existierender Pythonbefehl benennst, wirst du das an seiner Farbe erkennen! Dann solltest du dir am besten einen anderen Namen überlegen!

Wir können auch nur einzelne Funktionen eines Moduls importieren:

```python
from module import function_name
function_name()
```

Anstatt die Funktion `factorial` aus der obigen Aufgabe selbst zu schreiben, können wir nun einfach die gleichnamige Funktion aus dem Modul `math` verwenden:

```python
import math

math.factorial(5)
```

> Hinweis: Um herauszufinden welche Funktionen ein Modul zur Verfügung stellt, kannst du wieder die `<TAB>`-Vervollständigung im Jupyter Notebook verwenden:

```python
import math  # Importiere das Modul, indem du diese Zelle ausführst
```

```python
# math.<TAB> # Entferne das '#'-Symbol und drücke die <TAB>-Taste nach dem Punkt
```

> **Hinweis:** Diese Notation mit `name.funktions_name()` hast du auch schon bei den Objekten und deren Methoden gesehen (z.B. `dict.keys()`). Das sieht zwar gleich aus wie die Notation mit den Modulen, aber ist es ist nicht genau das Gleiche...Hier wird die Notation benutzt um klarzumachen aus welchem Modul wir eine Funktion benutzen. Also nicht verwirren lassen.

Die `from`-`import`-Syntax ist insbesondere für mathematische Ausdrücke hilfreich, sodass wir das Modul nicht immer schreiben müssen:

```python
from math import cos, sin
from math import pi as PI


def func(r, phi, theta):
    return r * cos(phi) * sin(theta)


func(1, 0, PI / 2)
```

## Module bieten vielseitige Funktionalität

Neben eingebauten Modulen wie `math` haben Python-Entwickler eine Vielzahl von Modulen für jeden Anwendungsbereich geschrieben. So können wir mit wenigen Zeilen Code äußerst komplexe Programme schreiben.

> Beispielsweise lesen wir mit `numpy` unseren Datensatz ein, berechnen mit `scipy` einen Fit und plotten beides mit `matplotlib`. Den Umgang mit diesen Modulen lernen wir im nächsten Kapitel.

Funktionen zur Berechnung von Mittelwert und Standardabweichung stellt bspw. `numpy` zur Verfügung:

```python
import numpy as np

li = [1, 2, 7, 3, 1, 3]
np.mean(li), np.std(li)
```

> Es gibt natürlich nicht nur Module für die wissenschaftliche Anwendung. Python wird höchst vielseitig eingesetzt, sodass du bspw. auch
> - eine [Webseite erstellen](http://getpelican.com),
> - einen [Webserver programmieren](http://www.djangoproject.com) oder
> - ein [Spiel entwickeln](http://www.pygame.org) kannst!

### Aufgabe 5 - Gaussian

Importiere die Funktionen `exp` aus `numpy` oder `math`.

Definiere eine Funktion mit dem Namen `gaussian`, welche die Argumente `x`, `mu`, `sigma` und `A` annimmt und den Wert $$A\cdot \,\mathrm{exp}\!\left(\frac{(x-\mu)^2}{2\cdot\sigma^2}\right)$$ zurückgibt.

<!-- grade_id: 103-3-sol, points: 0, grade: false, solution: true -->
```python
from numpy import exp, sqrt


def gaussian(x, mu, sigma, A):
    # DEINE LÖSUNG HIER
    pass  # `pass` ist ein Platzhalter und sagt Python, dass es einfach nix machen soll. Entferne es, wenn du eine Lösung hast
```

<!-- grade_id: 103-3-test, points: 1, grade: true, solution: false -->
```python
try:
    exp
except NameError:
    raise NameError(
        "Es gibt keine Funktion mit dem Namen 'exp'. Hast du die Funktion mit der `from module import function_name` Syntax importiert?"
    )

try:
    gaussian
except NameError:
    raise NameError(
        "Es gibt keine Funktion mit dem Namen 'gaussian'. Stelle sicher, dass deine Funktion so benannt ist."
    )

assert (
    gaussian(x=0, mu=0, sigma=1, A=1) == 1
), "gaussian(x=0, mu=0, sigma=1, A=1) sollte 1 ergeben. Prüfe den Code deiner Funktion."
assert (
    abs(gaussian(x=0, mu=1, sigma=2, A=3) - 3.4) < 0.1
), "gaussian(x=0, mu=1, sigma=2, A=3) sollte ca 3.4 ergeben. Prüfe den Code deiner Funktion."
print("👍 Stimmt so.")
```

---

Nun kannst du vollständige Programme schreiben und Funktionen aus Modulen verwenden. Erinnere dich daran - du musst nicht alles selbst schreiben! Baue lieber auf der Vorarbeit von schlauen Entwicklern auf der ganzen Welt auf, die schon hochoptimierte und getestete Lösungen für viele Probleme geschrieben haben.

In den nächsten drei Lektionen lernen wir die Grundlagen jeweils eines Moduls, das in der wissenschaftlichen Programmierung mit Python allgegenwärtig ist und beginnen mit dem Numerik-Modul _Numpy_.

[Startseite](index.ipynb) | [**>> 201 - Numerik mit Numpy**](201_Numerik_mit_Numpy.ipynb)

```python

```