Kurs 5.3 "Vom Gehirn Lernen" -- Python-Tutorial | [Startseite](index.ipynb)

---

# 202 - Wissenschaftliche Grafiken mit Matplotlib

**"Wissenschaftliche Grafiken" ist ein umständlicher und viel zu langer Begriff, deswegen hat sich das englische "Plot" dafür eingebürgert. Und in guter denglisch Manier hat man auch gleich noch das passende Verb "plotten" für das Erstellen der Grafiken gebastelt.**

- [Einfaches Plotten](#Einfaches-Plotten)
- [Plots gestalten](#Plots-gestalten)
- [Aufgabe 1 - Plot Design](#Aufgabe-1---Plot-Design)
- [Plots speichern](#Plots-speichern)
- [Aufgabe 2 - Temperaturen in Heidelberg visualisieren](#Aufgabe-2---Temperaturen-in-Heidelberg-visualisieren)

Wir können mit Numpy nun numerisch Daten auswerten, und um diese zu visualisieren können wir mit dem **Matplotlib** Modul Plots von wissenschaftlicher Qualität erstellen.

Per Konvention importieren wir das PyPlot Submodul von Matplotlib unter dem Namen `plt`:

```python
% matplotlib inline  # das ist einer sogenannter "cell magic command",
# also ein -- naja -- "magischer" Jupyter-Befehl, der dem Jupyter-Notebook mitteilt,
# dass die Plots direkt unter den Zellen angezeigt werden sollen.
# kleiner Hinweis: eigentlich ist dieser Befehl heutzutage obsolet, aber auf
# EBRAINS wird er scheinbar doch noch benötigt...

import matplotlib.pyplot as plt  # Die Abkürzung `plt` ist Konvention

# Numpy brauchen wir immer.
import numpy as np
```

## Einfaches Plotten

PyPlot stellt einige grundlegende Funktionen zur Verfügung, mit denen wir Daten schnell plotten können:

- **`plot`** nimmt x- und y-Daten sowie eine Vielzahl von optionalen Argumenten zur Konfiguration an. Fehlen die x-Daten, werden die Indizes der y-Daten verwendet:

```python
x = np.arange(-10, 10, 0.1)
y = np.arange(-10, 10, 0.1) ** 2
plt.plot(x, y)
```

```python
plt.plot(y)  # Beachte die x-Achse die du hier bekommst
```

Entfernt das '`#`'-Zeichen und schaut euch die Argumente in der Dokumentation an, mit denen wir den Plot konfigurieren können:

```python
# plt.plot?
```

Matplotlib bietet unzählige Möglichkeiten, das Aussehen von Plots anzupassen! Ihr könnt oben einige ausprobieren, wie bspw. `color`, `linestyle` und `linewidth`.

- **`scatter`** plottet Datenpunkte statt einer Linie:

```python
x = np.arange(10)
y = np.arange(10) ** 2
plt.scatter(x, y)
```

```python
# plt.scatter?
```

- **`errorbar`** nimmt zusätzlich Fehler in x- und y-Richtung an und zeichnet Fehlerbalken:

```python
x = np.arange(10)
y = np.arange(10) ** 2
err = np.arange(1, 6, 0.5)

plt.errorbar(x, y, yerr=err)
```

```python
# plt.errorbar?
```

- **`hist`** plottet schnell ein Histogramm:

```python
_ = plt.hist(np.random.normal(size=100), bins=10)
```

```python
# plt.hist?
```

> **Hinweis:** Alternativ könnt ihr Histogramme mit `numpy.histogram` und `matplotlib.pyplot.bar` plotten. Diese Methode bietet etwas mehr Kontrolle über die Berechnung der Histogrammdaten.

- **`imshow`** visualisiert zweidimensionale Daten:

```python
random_values = np.random.random((64, 64))
plt.imshow(random_values, interpolation="none")
plt.colorbar()  # die Farbe codiert den Wert der Einträge des zweidimensionalen arrays
```

```python
# plt.imshow?
```

## Plots gestalten

Mit Matplotlib könnt ihr jegliche Plots erstellen, die ihr euch vorstellen könnt. Überlegt euch zuerst, wie der Plot aussehen soll. Die [Gallerie](http://matplotlib.org/gallery.html) kann dabei helfen. Mit einer beherzten Websuche, die meist direkt zur [Dokumentation](http://matplotlib.org/api/pyplot_api.html) führt, findet ihr dann die Funktionen, die ihr braucht.

> **Hinweis:** Plotten ist eine Kunst.

> **Weiterer Hinweis:** Versucht's mal mit `plt.xkcd()`.

### Titel, und Achsen und Legende

Wir können Strings benutzen um die Plots mit Titel und Achsenbeschriftungen zu versehen.

Üblicherweise beschriftet man eine Achse mit dem Namen der Variablen und (falls möglich) mit der physikalischen Einheit der Größe in eckigen Klammern (z.B. `[cm]`).

```python
x = np.linspace(0, 2 * np.pi, 100)
y1 = np.sin(x)
y2 = np.cos(x)

# Plot mit Label für Legende
plt.plot(x, y1, label="sin(x)")
plt.plot(x, y2, label="cos(x)")

# Titel
plt.title("Oszillation")

# Achsenlimits
plt.xlim(0, 2 * np.pi)
plt.ylim(-1, 1)

# Achsenbeschriftungen
plt.xlabel("Winkel x [rad]")
plt.ylabel("Auslenkung y [cm]")

# Legende
plt.legend(loc="lower left")
```

### Aufgabe 1 - Plot Design

a) Erstelle einen Plot für y und z

b) Beschrifte die y-Kurve mit "y = f(x)" und die z-Kurve mit "z = g(x)" in einer Legende

c) Ändere die Farbe von y auf grün und von z auf rot.

d) Verschiebe die Legende nach oben und in die Mitte des Plots

e) Beschrifte die Achsen mit "x" und "y/z"

f) Ändere den Bereich der y-Achse des Plots zu Werten zwischen -30 und 110

g) Markiere jeden geplotteten Datenpunkt in den plots (Stichwort "marker") mit einem Punkt für y und einem Stern für z

h) Ändere den linestyle für y zu gestrichelt

```python
x = np.linspace(-10, 10, 16)
y = 3 * x + 10
z = x**2

# DEIN CODE HIER
```

## Plots speichern

Mit `matplotlib.pyplot.savefig` könnt ihr einen Plot als Bilddatei speichern:

```python
x = np.linspace(-10, 10, 16)
y = 3 * x + 10
plt.plot(x, y)
plt.savefig("plots/my_plot.png")
```

> **Hinweis:** Du kannst die Größe (in inches) und Auflösung (in dots-per-inch dpi) des gespeicherten Plots anpassen:
>
> ```python
> plt.savefig(filename, dpi=150, filesize=(10, 5))
> ```

### Aufgabe 2 - Temperaturen in Heidelberg visualisieren

Wir visualisieren nun die Temperaturdaten mit denen wir schon numerisch gearbeitet haben. Lest die Daten zunächst erneut ein:

```python
# Zelle ausführen, um Daten einzulesen
data = np.loadtxt("data/temperatures.txt")
date, T = data[np.abs(data[:, 1]) != 99, :].transpose()
```

a) Plottet den Temperaturverlauf über den gesamten Zeitraum. Vergesst nicht Titel, Achsenbeschriftungen und Legende.

<!-- grade_id: 202-1a-sol, points: 0, grade: false, solution: true -->
```python
# DEIN CODE HIER
```

Sieht dein Plot etwa so aus?

![Temperaturverlauf in Heidelberg](plots/202-1a.png)

```python
# Setze die Variable auf `True` wenn du mit deinem Plot zufrieden bist:
plot_1a_fertig = False
```

<!-- grade_id: 202-1a-test, points: 1, grade: true, solution: false -->
```python
assert (
    plot_1a_fertig
), "Versuche dich an Aufgabe 1a, bis du mit deinem Plot zufrieden bist."
print("Sieht gut aus.")
```

b) Plottet den kombinierten Jahres-Temperaturverlauf, also die Temperaturdaten aufgetragen über dem Zeitpunkt innerhalb des Jahres.

Plottet nur die Datenpunkte und keine verbindende Linie. Vergesst nicht Titel, Achsenbeschriftungen und Legende.

**Hinweis:** Den Zeitpunkt innerhalb eines Jahres erhaltet ihr wieder mit dem Modulo Operator: `date % 1`

<!-- grade_id: 202-1b-sol, points: 0, grade: false, solution: true -->
```python
# DEIN CODE HIER
```

Sieht dein Plot etwa so aus?

![Jahres-Temperaturverlauf in Heidelberg](plots/202-1b.png)

```python
# Setze die Variable auf `True` wenn du mit deinem Plot zufrieden bist:
plot_1b_fertig = False
```

<!-- grade_id: 202-1b-test, points: 1, grade: true, solution: false -->
```python
assert (
    plot_1b_fertig
), "Versuche dich an Aufgabe 1b, bis du mit deinem Plot zufrieden bist."
print("🖼 Wunderschön.")
```

---

Nun kannst du Daten mit Numpy analysieren und mit Matplotlib plotten. 

[Startseite](index.ipynb) | [**>> 301 - Abschlussaufgabe**](301_Gelateria.ipynb)

```python

```